## Exercise 01.1

Degrees Fahrenheit ($T_F$) are converted to degrees Celsius ($T_c$) using the formula:

$$
T_c = 5(T_f - 32)/9
$$

Write a program to convert 78 degrees Fahrenheit to degrees Celsius and print the result.
Write your program such that you can easily change the temperature in Fahrenheit that you are converting from.

Use the variable name `T_c` for the temperature in degrees Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python
T_f = 78
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
assert round(T_c - 25.55555555555555, 10) == 0.0
```

## Exercise 01.2

WOrking with NASA, you have been tasked with developing a tool that predicts the trajectory of orbitals. A component is to compute the speed of objects in free fall for a range of scenarios.

The impact speed of an object depends on the height on the launched object, the initial velocity and the acceleration due to gravity. 
Write a program that computes the impact speed, with variables for:

1. Height (meter)
1. Initial Velocity (meter/second)
1. Acceleration (meter/second^2)

Test your program with a height of 10m, upward velocity of 5m/s, and an acceleration due to gravity of 9.8 m/s^2.
Use the variable name `impactSpeed` for the speed upon impact.

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
```python
# Assign data to variables
acceleration = 9.8 

# YOUR CODE HERE
raise NotImplementedError()
```

```python
assert round(impactSpeed - 14.866, 2) == 0.0
```